<?xml version="1.0" encoding="UTF-8" ?>
        <quiz>
	<question type="category">
		<category><text>Stat110-115/B6/Hypothesis test for the difference between two means/210119130233-V1-A</text></category>
		<info format="html"><text></text></info>
		<idnumber></idnumber>
	</question>
	
		<question type="description">
		<name>
		  <text>210119130233 Introduction</text>
		</name>
		<questiontext format="html">
		  <text><![CDATA[<p dir="ltr" style="text-align: left;">This exercise is designed to be used with the statistical package R within RStudio. RStudio is on the lab computers \(MA124\) or you can download it to your own computer.
        The data file for this question is a sample from a built-in R data set\, ##iris##. The data set contains the measurements\, in centimeters\, of the variables ##Sepal_length##\, ##Sepal_width##\, ##Petal_length## and ##Petal_width##\, respectively\, for flowers from each of 3 species of iris. The ##Species## variable has the factor levels; Iris setosa \(##1##\)\, versicolor \(##2##\)\, and virginica \(##3##\).
        Using the link below\, save a text version \(file extension .txt\) of the datafile to a folder on your computer. Then open RStudio\, set your working directory to the folder where you saved the data\, and load the data using the ##read.table\(\)## function. Remember to assign the data to an R object \(i.e. give your dataset a name e.g. ##mydata##\). Note: this datafile has a header in it\, so you should set ##header=TRUE##. Ensure you set the working directory to where you saved the file in order to read it in\, using double "\\" and not single "\" in the filepath name. Further instructions on using R and loading in data can be found in the Introduction to R document on the resources page and in the course slides.
        #datafileT
        Once you have loaded in the data\, use the command ##head\(mydata\)##. This allows you to see the variable names and the format of the data.
        
        
        
        HAS DATA FILE ==> irishtest;400

        
        Data from datafile is: 

            <table>
            <tr><td>Sepal_Length</td><td>Sepal_Width</td><td>Petal_Length</td><td>Petal_Width</td><td>Species</td></tr>
                    <tr><td>6</td><td>3.4</td><td>4.5</td><td>1.6</td><td>2</td></tr>
                    <tr><td>5.7</td><td>2.5</td><td>5</td><td>2</td><td>3</td></tr>
                    <tr><td>5.8</td><td>2.7</td><td>4.1</td><td>1</td><td>2</td></tr>
                    <tr><td>5</td><td>3.6</td><td>1.4</td><td>0.2</td><td>1</td></tr>
                    <tr><td>5</td><td>3.5</td><td>1.6</td><td>0.6</td><td>1</td></tr>
                    <tr><td>6.2</td><td>2.8</td><td>4.8</td><td>1.8</td><td>3</td></tr>
                    <tr><td>7.3</td><td>2.9</td><td>6.3</td><td>1.8</td><td>3</td></tr>
                    <tr><td>6.4</td><td>3.2</td><td>4.5</td><td>1.5</td><td>2</td></tr>
                    <tr><td>5.1</td><td>2.5</td><td>3</td><td>1.1</td><td>2</td></tr>
                    <tr><td>5.4</td><td>3</td><td>4.5</td><td>1.5</td><td>2</td></tr>
                    <tr><td>5.8</td><td>2.7</td><td>5.1</td><td>1.9</td><td>3</td></tr>
                    <tr><td>6.4</td><td>2.7</td><td>5.3</td><td>1.9</td><td>3</td></tr>
                    <tr><td>6.9</td><td>3.1</td><td>5.4</td><td>2.1</td><td>3</td></tr>
                    <tr><td>5.2</td><td>3.5</td><td>1.5</td><td>0.2</td><td>1</td></tr>
                    <tr><td>6.1</td><td>3</td><td>4.9</td><td>1.8</td><td>3</td></tr>
                    <tr><td>4.8</td><td>3.1</td><td>1.6</td><td>0.2</td><td>1</td></tr>
                    <tr><td>6.1</td><td>2.6</td><td>5.6</td><td>1.4</td><td>3</td></tr>
                    <tr><td>5.6</td><td>2.9</td><td>3.6</td><td>1.3</td><td>2</td></tr>
                    <tr><td>5.7</td><td>2.6</td><td>3.5</td><td>1</td><td>2</td></tr>
                    <tr><td>4.9</td><td>2.5</td><td>4.5</td><td>1.7</td><td>3</td></tr>
                    <tr><td>7.4</td><td>2.8</td><td>6.1</td><td>1.9</td><td>3</td></tr>
                    <tr><td>6.7</td><td>3.3</td><td>5.7</td><td>2.5</td><td>3</td></tr>
                    <tr><td>5.5</td><td>2.6</td><td>4.4</td><td>1.2</td><td>2</td></tr>
                    <tr><td>5.7</td><td>4.4</td><td>1.5</td><td>0.4</td><td>1</td></tr>
                    <tr><td>5</td><td>2</td><td>3.5</td><td>1</td><td>2</td></tr>
                    <tr><td>6</td><td>2.2</td><td>5</td><td>1.5</td><td>3</td></tr>
                    <tr><td>6.9</td><td>3.1</td><td>4.9</td><td>1.5</td><td>2</td></tr>
                    <tr><td>5.1</td><td>3.5</td><td>1.4</td><td>0.2</td><td>1</td></tr>
                    <tr><td>6.1</td><td>2.9</td><td>4.7</td><td>1.4</td><td>2</td></tr>
                    <tr><td>6.3</td><td>2.8</td><td>5.1</td><td>1.5</td><td>3</td></tr>
                    <tr><td>6.7</td><td>3.1</td><td>4.4</td><td>1.4</td><td>2</td></tr>
                    <tr><td>5.4</td><td>3.9</td><td>1.3</td><td>0.4</td><td>1</td></tr>
                    <tr><td>4.8</td><td>3</td><td>1.4</td><td>0.3</td><td>1</td></tr>
                    <tr><td>4.7</td><td>3.2</td><td>1.3</td><td>0.2</td><td>1</td></tr>
                    <tr><td>5</td><td>3.2</td><td>1.2</td><td>0.2</td><td>1</td></tr>
                    <tr><td>5</td><td>3.3</td><td>1.4</td><td>0.2</td><td>1</td></tr>
                    <tr><td>6.5</td><td>3</td><td>5.2</td><td>2</td><td>3</td></tr>
                    <tr><td>5</td><td>3</td><td>1.6</td><td>0.2</td><td>1</td></tr>
                    <tr><td>5.6</td><td>2.7</td><td>4.2</td><td>1.3</td><td>2</td></tr>
                    <tr><td>4.3</td><td>3</td><td>1.1</td><td>0.1</td><td>1</td></tr>
                    <tr><td>4.6</td><td>3.6</td><td>1</td><td>0.2</td><td>1</td></tr>
                    <tr><td>4.4</td><td>3.2</td><td>1.3</td><td>0.2</td><td>1</td></tr>
                    <tr><td>6.6</td><td>2.9</td><td>4.6</td><td>1.3</td><td>2</td></tr>
                    <tr><td>4.9</td><td>2.4</td><td>3.3</td><td>1</td><td>2</td></tr>
                    <tr><td>7</td><td>3.2</td><td>4.7</td><td>1.4</td><td>2</td></tr>
                    <tr><td>6.7</td><td>3</td><td>5.2</td><td>2.3</td><td>3</td></tr>
                    <tr><td>4.6</td><td>3.1</td><td>1.5</td><td>0.2</td><td>1</td></tr>
                    <tr><td>5.9</td><td>3.2</td><td>4.8</td><td>1.8</td><td>2</td></tr>
                    <tr><td>6.8</td><td>2.8</td><td>4.8</td><td>1.4</td><td>2</td></tr>
                    <tr><td>4.4</td><td>3</td><td>1.3</td><td>0.2</td><td>1</td></tr>
                    </table></p>]]></text>
		</questiontext>
		<generalfeedback format="html">
		  <text><![CDATA[<p dir="ltr" style="text-align: left;"></p>]]></text>
		</generalfeedback>
		<defaultgrade>0</defaultgrade>
		<penalty>0</penalty>
		<hidden>0</hidden>
		<idnumber></idnumber>
	  </question>
	  

	<question type="multichoice">
    <name>
      <text>210119130233 Multi-Choice Question - Q1</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p dir="ltr" style="text-align: left;">~Use the commands <b> virginica<-mydata\$Sepal_Length{@mydata\$Species==3@}  </b> and <b>setosa<-mydata\$Sepal_Length{@mydata\$Species==1@}  </b> to store the sepal lengths of the Virginica and Setosa species. In this question you will compare the sample means to estimate the true difference in sepal length means of these species.
        If $V$ represents “Virginica” and $S$ represents “Setosa”\, which one of the following is the appropriate pair of hypotheses?
        
        
	
	==> MULTI-CHOICE QUESTION HAS VARIABLES, LISTED HERE: 
	
	<questionvariables>
	<text>{@xyz@}:0;
{@variables@}:{@xyz@};
{@MeanSL@}:5.7;
{@size@}:50;
{@VarSL@}:0.689795918367347;
{@MeanSL_Se@}:4.88235294117647;
{@MeanSL_Vi@}:6.33333333333333;
{@MeanSL_Ve@}:5.94444444444444;
{@VarSL_Se@}:0.135294117647059;
{@VarSL_Vi@}:0.405238095238095;
{@VarSL_Ve@}:0.441437908496732;
{@Length_Se@}:17;
{@Length_Vi@}:15;
{@Length_Ve@}:18;
{@DF_WS@}:21.807413040466;
{@variables@}:{@MeanSL@}\,{@size@}\,{@VarSL@}\,{@MeanSL_Se@}\,{@MeanSL_Vi@}\,{@MeanSL_Ve@}\,{@VarSL_Se@}\,{@VarSL_Vi@}\,{@VarSL_Ve@}\,{@Length_Se@}\,{@Length_Vi@}\,{@Length_Ve@}\,{@DF_WS@};
</text>
	</questionvariables></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p dir="ltr" style="text-align: left;">Well done</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <single>true</single>
    <shuffleanswers>true</shuffleanswers>
    <answernumbering>abc</answernumbering>
    <showstandardinstruction>0</showstandardinstruction>
    <correctfeedback format="html">
      <text>Your answer is correct.</text>
    </correctfeedback>
    <partiallycorrectfeedback format="html">
      <text>Your answer is partially correct.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="html">
      <text>Your answer is incorrect.</text>
    </incorrectfeedback>
    <shownumcorrect/>
		<answer fraction="0" format="html">
		<text><![CDATA[<p dir="ltr" style="text-align: left;">$H_0: \mu_V - \mu_S \ne 0; \;\;  H_A: \mu_V - \mu_S = 0$</p>]]></text>
		<feedback format="html">
		  <text><![CDATA[<p dir="ltr" style="text-align: left;"></p>]]></text>
		</feedback>
	  </answer>
		<answer fraction="0" format="html">
		<text><![CDATA[<p dir="ltr" style="text-align: left;">$H_0: \bar{x}_V - \bar{x}_S =0; \;\;  H_A: \bar{x}_V- \bar{x}_S \ne 0$</p>]]></text>
		<feedback format="html">
		  <text><![CDATA[<p dir="ltr" style="text-align: left;"></p>]]></text>
		</feedback>
	  </answer>
		<answer fraction="0" format="html">
		<text><![CDATA[<p dir="ltr" style="text-align: left;">$H_0: \bar{x}_V = \bar{x}_S ; \;\;  H_A: \bar{x}_V \ne \bar{x}_S$</p>]]></text>
		<feedback format="html">
		  <text><![CDATA[<p dir="ltr" style="text-align: left;"></p>]]></text>
		</feedback>
	  </answer>
		<answer fraction="0" format="html">
		<text><![CDATA[<p dir="ltr" style="text-align: left;">$H_0: \mu_V ≠ \mu_S ; \;\;  H_A: \mu_V = \mu_S$</p>]]></text>
		<feedback format="html">
		  <text><![CDATA[<p dir="ltr" style="text-align: left;"></p>]]></text>
		</feedback>
	  </answer>
		<answer fraction="100" format="html">
		<text><![CDATA[<p dir="ltr" style="text-align: left;">$H_0: \mu_V = \mu_S ; \;\;  H_A: \mu_V ≠ \mu_S$</p>]]></text>
		<feedback format="html">
		  <text><![CDATA[<p dir="ltr" style="text-align: left;"></p>]]></text>
		</feedback>
	  </answer>
	</question>
	
	<question type="stack">
    <name>
      <text>210119130233 Stack Question - Q2</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[
			  <p>What is the standard error of the estimated difference in means? Hint: To calculate this\, it might be helpful to use the commands length\(virginica\) and length\(setosa\) to know how many virginica and setosa observations there are in this sample.
        
        </p>
			  <p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html"><text></text></generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion><text>2023060500</text></stackversion>	

	
	<questionvariables>
	<text>{@xyz@}:0;
{@variables@}:{@xyz@};
{@MeanSL@}:5.7;
{@size@}:50;
{@VarSL@}:0.689795918367347;
{@MeanSL_Se@}:4.88235294117647;
{@MeanSL_Vi@}:6.33333333333333;
{@MeanSL_Ve@}:5.94444444444444;
{@VarSL_Se@}:0.135294117647059;
{@VarSL_Vi@}:0.405238095238095;
{@VarSL_Ve@}:0.441437908496732;
{@Length_Se@}:17;
{@Length_Vi@}:15;
{@Length_Ve@}:18;
{@DF_WS@}:21.807413040466;
{@variables@}:{@MeanSL@}\,{@size@}\,{@VarSL@}\,{@MeanSL_Se@}\,{@MeanSL_Vi@}\,{@MeanSL_Ve@}\,{@VarSL_Se@}\,{@VarSL_Vi@}\,{@VarSL_Ve@}\,{@Length_Se@}\,{@Length_Vi@}\,{@Length_Ve@}\,{@DF_WS@};
</text>
	</questionvariables>
	<specificfeedback format="html"><text></text></specificfeedback>
	<questionnote><text></text></questionnote>
	<questiondescription format="html"><text></text></questiondescription>
	<questionsimplify>1</questionsimplify>
	<assumepositive>0</assumepositive>
	<assumereal>0</assumereal>
	<prtcorrect format="html"><text>Correct answer, well done.</text></prtcorrect>
	<prtpartiallycorrect format="html"><text>Your answer is partially correct.</text></prtpartiallycorrect>
	<prtincorrect format="html"><text>Incorrect answer.</text></prtincorrect>
	<multiplicationsign>cross</multiplicationsign>
	<sqrtsign>1</sqrtsign>
	<complexno>i</complexno>
	<inversetrig>cos-1</inversetrig>
	<logicsymbol>lang</logicsymbol>
	<matrixparens>[</matrixparens>
	<variantsselectionseed></variantsselectionseed>
		<input>
		  <name>ans1</name>
		  <type>algebraic</type>
		  <tans>pow\({@VarSL_Se@}/{@Length_Se@}+{@VarSL_Vi@}/{@Length_Vi@}\,0.5\)</tans>
		  <boxsize>15</boxsize>
		  <strictsyntax>1</strictsyntax>
		  <insertstars>0</insertstars>
		  <syntaxhint></syntaxhint>
		  <syntaxattribute>0</syntaxattribute>
		  <forbidwords></forbidwords>
		  <allowwords></allowwords>
		  <forbidfloat>0</forbidfloat>
		  <requirelowestterms>0</requirelowestterms>
		  <checkanswertype>1</checkanswertype>
		  <mustverify>1</mustverify>
		  <showvalidation>1</showvalidation>
		  <options></options>
		</input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>The standard error is $\sqrt{\frac{s_S^2}{n_S}+\frac{s_V^2}{n_V}}=\sqrt{\frac{{@{@decimalplacesDP\({@VarSL_Se@}\,4\)@}@}}{{@Length_Se@}}+\frac{{@{@decimalplacesDP\({@VarSL_Vi@}\,4\)@}@}}{{@Length_Vi@}}}$\nwhere $s_S$ and $s_V$ are the sample standard deviations for setosa and virginica species\, and $n_S$ and $n_V$ are the number of setosa and virginica observations in the sample.</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>pow\({@VarSL_Se@}/{@Length_Se@}+{@VarSL_Vi@}/{@Length_Vi@}\,0.5\)</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">Correct answer</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">Incorrect answer</p>]]></text>
        </falsefeedback>
      </node>
    </prt><qtest>
	<testcase>1</testcase>
	<description></description>
		<testinput><name>ans2</name><value></value></testinput>
	
		<expected>
		  <name>prt2</name>
		  <expectedscore>0.0000000</expectedscore>
		  <expectedpenalty>0.1000000</expectedpenalty>
		  <expectedanswernote>prt2-1-F</expectedanswernote>
		</expected></qtest></question>
	
	<question type="stack">
    <name>
      <text>210119130233 Stack Question - Q3</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[
			  <p>Using the order of differencing $\(\bar{x}_V-\bar{x}_S\)$\, calculate the test statistic for this test.
        
        </p>
			  <p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html"><text></text></generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion><text>2023060500</text></stackversion>	

	
	<questionvariables>
	<text>{@xyz@}:0;
{@variables@}:{@xyz@};
{@MeanSL@}:5.7;
{@size@}:50;
{@VarSL@}:0.689795918367347;
{@MeanSL_Se@}:4.88235294117647;
{@MeanSL_Vi@}:6.33333333333333;
{@MeanSL_Ve@}:5.94444444444444;
{@VarSL_Se@}:0.135294117647059;
{@VarSL_Vi@}:0.405238095238095;
{@VarSL_Ve@}:0.441437908496732;
{@Length_Se@}:17;
{@Length_Vi@}:15;
{@Length_Ve@}:18;
{@DF_WS@}:21.807413040466;
{@variables@}:{@MeanSL@}\,{@size@}\,{@VarSL@}\,{@MeanSL_Se@}\,{@MeanSL_Vi@}\,{@MeanSL_Ve@}\,{@VarSL_Se@}\,{@VarSL_Vi@}\,{@VarSL_Ve@}\,{@Length_Se@}\,{@Length_Vi@}\,{@Length_Ve@}\,{@DF_WS@};
</text>
	</questionvariables>
	<specificfeedback format="html"><text></text></specificfeedback>
	<questionnote><text></text></questionnote>
	<questiondescription format="html"><text></text></questiondescription>
	<questionsimplify>1</questionsimplify>
	<assumepositive>0</assumepositive>
	<assumereal>0</assumereal>
	<prtcorrect format="html"><text>Correct answer, well done.</text></prtcorrect>
	<prtpartiallycorrect format="html"><text>Your answer is partially correct.</text></prtpartiallycorrect>
	<prtincorrect format="html"><text>Incorrect answer.</text></prtincorrect>
	<multiplicationsign>cross</multiplicationsign>
	<sqrtsign>1</sqrtsign>
	<complexno>i</complexno>
	<inversetrig>cos-1</inversetrig>
	<logicsymbol>lang</logicsymbol>
	<matrixparens>[</matrixparens>
	<variantsselectionseed></variantsselectionseed>
		<input>
		  <name>ans1</name>
		  <type>algebraic</type>
		  <tans>\({@MeanSL_Vi@}-{@MeanSL_Se@}\)/U{@-1@}</tans>
		  <boxsize>15</boxsize>
		  <strictsyntax>1</strictsyntax>
		  <insertstars>0</insertstars>
		  <syntaxhint></syntaxhint>
		  <syntaxattribute>0</syntaxattribute>
		  <forbidwords></forbidwords>
		  <allowwords></allowwords>
		  <forbidfloat>0</forbidfloat>
		  <requirelowestterms>0</requirelowestterms>
		  <checkanswertype>1</checkanswertype>
		  <mustverify>1</mustverify>
		  <showvalidation>1</showvalidation>
		  <options></options>
		</input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>$\frac{\(\bar{x}_V-\bar{x}_S\)-0}{SE}=\frac{{@{@decimalplacesDP\({@MeanSL_Vi@}\,4\)@}@}-{@{@decimalplacesDP\({@MeanSL_Se@}\,4\)@}@}}{{@{@decimalplacesDP\(pow\({@VarSL_Se@}/{@Length_Se@}+{@VarSL_Vi@}/{@Length_Vi@}\,0.5\)\,4\)@}@}}={@{@decimalplacesDP\(\({@MeanSL_Vi@}-{@MeanSL_Se@}\)/\(sqrt\({@VarSL_Se@}/{@Length_Se@}+{@VarSL_Vi@}/{@Length_Vi@}\)\)\,3\)@}@}$</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>\({@MeanSL_Vi@}-{@MeanSL_Se@}\)/U{@-1@}</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">Correct answer</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">Incorrect answer</p>]]></text>
        </falsefeedback>
      </node>
    </prt><qtest>
	<testcase>1</testcase>
	<description></description>
		<testinput><name>ans2</name><value></value></testinput>
	
		<expected>
		  <name>prt2</name>
		  <expectedscore>0.0000000</expectedscore>
		  <expectedpenalty>0.1000000</expectedpenalty>
		  <expectedanswernote>prt2-1-F</expectedanswernote>
		</expected></qtest></question>
	
	<question type="stack">
    <name>
      <text>210119130233 Stack Question - Q4</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[
			  <p>Using the estimated degrees of freedom\, $\nu={@{@decimalplacesDP\({@DF_WS@}\,2\)@}@}$\, the $p$-value is:
        
        </p>
			  <p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html"><text></text></generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion><text>2023060500</text></stackversion>	

	
	<questionvariables>
	<text>{@xyz@}:0;
{@variables@}:{@xyz@};
{@MeanSL@}:5.7;
{@size@}:50;
{@VarSL@}:0.689795918367347;
{@MeanSL_Se@}:4.88235294117647;
{@MeanSL_Vi@}:6.33333333333333;
{@MeanSL_Ve@}:5.94444444444444;
{@VarSL_Se@}:0.135294117647059;
{@VarSL_Vi@}:0.405238095238095;
{@VarSL_Ve@}:0.441437908496732;
{@Length_Se@}:17;
{@Length_Vi@}:15;
{@Length_Ve@}:18;
{@DF_WS@}:21.807413040466;
{@variables@}:{@MeanSL@}\,{@size@}\,{@VarSL@}\,{@MeanSL_Se@}\,{@MeanSL_Vi@}\,{@MeanSL_Ve@}\,{@VarSL_Se@}\,{@VarSL_Vi@}\,{@VarSL_Ve@}\,{@Length_Se@}\,{@Length_Vi@}\,{@Length_Ve@}\,{@DF_WS@};
</text>
	</questionvariables>
	<specificfeedback format="html"><text></text></specificfeedback>
	<questionnote><text></text></questionnote>
	<questiondescription format="html"><text></text></questiondescription>
	<questionsimplify>1</questionsimplify>
	<assumepositive>0</assumepositive>
	<assumereal>0</assumereal>
	<prtcorrect format="html"><text>Correct answer, well done.</text></prtcorrect>
	<prtpartiallycorrect format="html"><text>Your answer is partially correct.</text></prtpartiallycorrect>
	<prtincorrect format="html"><text>Incorrect answer.</text></prtincorrect>
	<multiplicationsign>cross</multiplicationsign>
	<sqrtsign>1</sqrtsign>
	<complexno>i</complexno>
	<inversetrig>cos-1</inversetrig>
	<logicsymbol>lang</logicsymbol>
	<matrixparens>[</matrixparens>
	<variantsselectionseed></variantsselectionseed>
		<input>
		  <name>ans1</name>
		  <type>algebraic</type>
		  <tans>2*\(1-CTDist\(abs\(U{@-1@}\)\,{@DF_WS@}\)\)</tans>
		  <boxsize>15</boxsize>
		  <strictsyntax>1</strictsyntax>
		  <insertstars>0</insertstars>
		  <syntaxhint></syntaxhint>
		  <syntaxattribute>0</syntaxattribute>
		  <forbidwords></forbidwords>
		  <allowwords></allowwords>
		  <forbidfloat>0</forbidfloat>
		  <requirelowestterms>0</requirelowestterms>
		  <checkanswertype>1</checkanswertype>
		  <mustverify>1</mustverify>
		  <showvalidation>1</showvalidation>
		  <options></options>
		</input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><b>2*\(1-pt\({@{@decimalplacesDP\(\({@MeanSL_Vi@}-{@MeanSL_Se@}\)/\(sqrt\({@VarSL_Se@}/{@Length_Se@}+{@VarSL_Vi@}/{@Length_Vi@}\)\)\,3\)@}@}\,{@{@decimalplacesDP\({@DF_WS@}\,2\)@}@}\)\)</b> OR <b>2*pt\({@{@decimalplacesDP\(\({@MeanSL_Vi@}-{@MeanSL_Se@}\)/\(sqrt\({@VarSL_Se@}/{@Length_Se@}+{@VarSL_Vi@}/{@Length_Vi@}\)\)\,3\)@}@}\,{@{@decimalplacesDP\({@DF_WS@}\,2\)@}@}\,lower.tail=FALSE\)</b>\n\n<b>General method for finding $p$-values using a $t$-distribution:</b> Sketch the t-distribution and mark the mean \(the mean of the t-distribution is always zero\) add your test statistic. The $p$-value is the probability under the t-distribution of your test statistic $\textit{or a value more extreme}$\, ie away from the mean. You then need to double this probability to include the other tail \(this is because prior to collecting your data it is possible that deviations from the mean could occur in either direction from the mean\).\n\nIf your test statistic \(TS\) is $negative$ use:\n<b>2*pt\(TS\,$\nu$\,lower.tail=TRUE\)</b>\nequivalently\n<b>2*pt\(TS\,$\nu$\)</b>\nbecause lower.tail=TRUE is default and so does not need to be provided\nIf your test statistic is $positive$ use:\n<b>2*pt\(TS\,$\nu$\,lower.tail=FALSE\)</b> or  <b>2*\(1-pt\(TS\,$\nu$\)</b></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>2*\(1-CTDist\(abs\(U{@-1@}\)\,{@DF_WS@}\)\)</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">Correct answer</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">Incorrect answer</p>]]></text>
        </falsefeedback>
      </node>
    </prt><qtest>
	<testcase>1</testcase>
	<description></description>
		<testinput><name>ans2</name><value></value></testinput>
	
		<expected>
		  <name>prt2</name>
		  <expectedscore>0.0000000</expectedscore>
		  <expectedpenalty>0.1000000</expectedpenalty>
		  <expectedanswernote>prt2-1-F</expectedanswernote>
		</expected></qtest></question>
	

	<question type="multichoice">
    <name>
      <text>210119130233 Multi-Choice Question - Q5</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p dir="ltr" style="text-align: left;">Assuming the necessary assumptions are satisfied\, select the appropriate conclusion given the $p$-value you have calculated.
        
        
	
	==> MULTI-CHOICE QUESTION HAS VARIABLES, LISTED HERE: 
	
	<questionvariables>
	<text>{@xyz@}:0;
{@variables@}:{@xyz@};
{@MeanSL@}:5.7;
{@size@}:50;
{@VarSL@}:0.689795918367347;
{@MeanSL_Se@}:4.88235294117647;
{@MeanSL_Vi@}:6.33333333333333;
{@MeanSL_Ve@}:5.94444444444444;
{@VarSL_Se@}:0.135294117647059;
{@VarSL_Vi@}:0.405238095238095;
{@VarSL_Ve@}:0.441437908496732;
{@Length_Se@}:17;
{@Length_Vi@}:15;
{@Length_Ve@}:18;
{@DF_WS@}:21.807413040466;
{@variables@}:{@MeanSL@}\,{@size@}\,{@VarSL@}\,{@MeanSL_Se@}\,{@MeanSL_Vi@}\,{@MeanSL_Ve@}\,{@VarSL_Se@}\,{@VarSL_Vi@}\,{@VarSL_Ve@}\,{@Length_Se@}\,{@Length_Vi@}\,{@Length_Ve@}\,{@DF_WS@};
</text>
	</questionvariables></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p dir="ltr" style="text-align: left;">Well done</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <single>true</single>
    <shuffleanswers>true</shuffleanswers>
    <answernumbering>abc</answernumbering>
    <showstandardinstruction>0</showstandardinstruction>
    <correctfeedback format="html">
      <text>Your answer is correct.</text>
    </correctfeedback>
    <partiallycorrectfeedback format="html">
      <text>Your answer is partially correct.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="html">
      <text>Your answer is incorrect.</text>
    </incorrectfeedback>
    <shownumcorrect/>
		<answer fraction="0" format="html">
		<text><![CDATA[<p dir="ltr" style="text-align: left;">The $p$-value is greater than or equal to 0.05.  We therefore do not have sufficient evidence to reject $H_0$. The difference between the sample means is consistent with what we might expect to see from random variation under the null hypothesis.</p>]]></text>
		<feedback format="html">
		  <text><![CDATA[<p dir="ltr" style="text-align: left;"></p>]]></text>
		</feedback>
	  </answer>
		<answer fraction="0" format="html">
		<text><![CDATA[<p dir="ltr" style="text-align: left;">The $p$-value is greater than or equal to 0.05. We therefore have evidence to reject $H_0$ and can say that there is evidence of a difference in sepal length between the Virginica and Setosa species at the 5% level.</p>]]></text>
		<feedback format="html">
		  <text><![CDATA[<p dir="ltr" style="text-align: left;"></p>]]></text>
		</feedback>
	  </answer>
		<answer fraction="0" format="html">
		<text><![CDATA[<p dir="ltr" style="text-align: left;"> The $p$-value is less than 0.05. We therefore have evidence to reject $H_0$ and can say that there is evidence of a difference in sepal length between the Virginica and Setosa species at the 5% level.</p>]]></text>
		<feedback format="html">
		  <text><![CDATA[<p dir="ltr" style="text-align: left;"></p>]]></text>
		</feedback>
	  </answer>
		<answer fraction="0" format="html">
		<text><![CDATA[<p dir="ltr" style="text-align: left;">Because the $p$-value is smaller than 0.95 there will not be any significant result from this analysis.</p>]]></text>
		<feedback format="html">
		  <text><![CDATA[<p dir="ltr" style="text-align: left;"></p>]]></text>
		</feedback>
	  </answer>
		<answer fraction="0" format="html">
		<text><![CDATA[<p dir="ltr" style="text-align: left;">The $p$-value is less than 0.05. We therefore do not have sufficient evidence to reject $H_0$. The difference between the sample means is consistent with what we might expect to see from random variation under the null hypothesis.</p>]]></text>
		<feedback format="html">
		  <text><![CDATA[<p dir="ltr" style="text-align: left;"></p>]]></text>
		</feedback>
	  </answer>
	</question>
	</quiz>